# Funktionen

In diesem Abschnitt gehen wir darauf ein wie man Funktionen deklariert, aufruft und anschließend schauen wir uns noch einige spezielle Eigenschaften von Funktionen in Julia an.

Inhalt:
1. Wie man Funktionen deklariert 
2. Duck-Typing in Julia
3. Verändernde vs. nicht-verändernde Funktionen
4. Einige höherwertige Funktionen

## Wie man Funktionen deklariert
Julia bietet uns einige verschiedene Möglichkeiten, eine Funktion zu schreiben. Die erste erfordert die Schlüsselwörter `function` und `end`.

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

```julia
sayhi("C-3PO")
```

```julia
function f(x)
    x^2
end
```

```julia
f(42)
```

Wir können diese Funktionen auch in einer Zeile definieren, indem wir die Schlüsselwörter `function` und `end` weglassen und stattdessen ein `=` Zeichen verwenden.

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
f2(x) = x^2
```

```julia
sayhi2("Chewbacca")
f2(42)
```

Die letzte möglichkeit ist dass wir die Funktionen anonym deklatrieren. 

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
f3 = x -> x^2
```

```julia
sayhi3("Chewbacca")
f3(42)
```

## Duck-typing in Julia
*"If it quacks like a duck, it's a duck."* <br><br>
Julia-Funktionen funktionieren einfach mit allen Eingaben, die "Sinn" ergeben. <br><br>
Zum Beispiel funktioniert `sayhi` auch mit einer Ganzzahl anstatt einem String...

```julia
# wir definieren die Zwei Funktionen aus dem letzten Abschnitt noch mal:
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
function f(x)
    x^2
end
```

```julia
sayhi(55595472)
```

Die Funktion `f` Funktioniert auch mit einer Matrix. 

```julia
A = rand(3, 3)
A
```

```julia
f(A)
```

`f` funktioniert auch mit dem Strin "hi" weil mit `*` Strings konkatiniert werden.

```julia
f("hi")
```

Allerdings wird `f`nicht mit einem Vektor funktionieren. Könnt ihr euch vorstellen warum? 

```julia
v = rand(3)
```

```julia
# This won't work
f(v)
```

Anders als `A^2`, das eindeutig definiert ist, ist die Bedeutung von `v^2` für einen Vektor, `v`, keine eindeutige algebraische Operation.

## Mutating vs. non-mutating Funktionen (Verändernde vs. nicht-verändernde Funktionen)

In Julia gibt es zwei Arten von Funktionen: Mutating und non-mutating. Funktionen denen ein `!`folgt verändern ihre Eingabe. Funktionen ohne `!` verändern ihre Eingabe nicht. Ihr kennt das Konzept eventuell auch unter unter den name call by value und call by reference. 

Zum Beispiel, schauen wir uns `sort` und `sort!` an.

```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort(v)` gibt eine sortierte Liste zurück, die die gleichen Elemente wie `v` enthält, aber `v` bleibt unverändert. Das entspricht also call by value. <br><br>

Auf der anderen Seite, wenn wir `sort!(v)` ausführen, werden die Inhalte von v innerhalb des Arrays `v` sortiert. Das entspricht also call by reference.

```julia
sort!(v)
```

```julia
v
```

Allerdings ist dies nur eine best practice und nicht eine Regel. Es hat keine bedeutung auf Compilerebene. Julia erlaubt es uns, die Eingabe einer Funktion zu verändern, auch wenn wir kein `!` verwenden. 
Das bedeutet wenn wir eine Funktion selbst implementieren sollten wir darauf achten ob wir die Eingabe verändern wollen oder nicht.

## Einige höherwertige Funktionen

### map

`map` ist eine "höherwertige" Funktion in Julia, die *eine Funktion* als eines ihrer Eingabeargumente nimmt.
`map` wendet dann diese Funktion auf jedes Element der Datenstruktur an, die Sie ihr übergeben. Zum Beispiel, wenn wir


```julia
map(f, [1, 2, 3])
```
ausführen, gibt uns `map` eine Liste zurück, die die Funktion `f` auf jedes Element von `[1, 2, 3]` angewendet hat. <br><br>
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Wir haben also jedes element in dem vektor `[1, 2, 3]` quadriert.

Alternativ können wit `map` eine anonyme Funktion anstatt einer benannten Funktion übergeben. Zum Beispiel:

```julia
map(x -> x^3, [1, 2, 3])
```

### broadcast
`broadcast` ist eine andere höherwertige Funktion in Julia, die *eine Funktion* als eines ihrer Eingabeargumente nimmt. `broadcast` wendet dann diese Funktion auf jedes Element der Datenstruktur an, die Sie ihr übergeben. Sie kann als eine allgemeinere Version von `map`betrachtet werden.

```julia
broadcast(f, [1, 2, 3])
```

Wir haben wieder `f`auf alle Elemente von `[1, 2, 3]` angewendet. Diesmal haben wir `f` "broadcasted"! 

Wir können `broadcast` auch mit `.`aufrufen indem wir es zwischen den Namen der Funktion und ihren Eingabeargumenten platzieren. Zum Beispiel,

```julia
broadcast(f, [1, 2, 3])
```
ist das gleiche wie
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Das ermöglicht uns funktionen wie `f` auf Vektoren anzuwenden. 

Das bedeutet natürlich auch, dass wir `f` auf Matrizen anwenden können. Allerdings sind die Ergebnisse für 
```julia
f(A)
```
und
```julia
f.(A)
```
für die Matrix `A` nicht gleich:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

Das ist das gleiche Verhalten das wir schon kennen. Für `A`, ist dieser aufruf:
```
f(A) = A^2 = A * A
``` 

Wenn wir nun `.`, also `broadcast` verwenden, dann wird `f` auf jedes Element von `A` angewendet. Das ist also

```julia
B = f.(A)
```

Das Ergrebnis ist eine Matrix, die die Quadrate aller Elemente von `A` enthält.

Dieses Punkt-Syntax für Broadcasting erlaubt es uns relativ komplexe Elementweise Ausdrücke zu schreiben, die sich oft einfacher als eine mathematische Notation lesen lassen. Zum Beispiel können wir schreiben:

```julia
A .+ 2 .* f.(A) ./ A
```

Anstatt, folgendes schreiben zu müssen

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

beide führen die gleiche Berechnung aus. 

## Übung: Funktionen

#### 1.1 
Schreibe eine Funktion `add_one` die 1 zu ihrer Eingabe hinzufügt.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 1.2 
Nutze `map` oder `broadcast` um jedes Element von `A` um 1 zu erhöhen und weise es einer Variable `A1` zu.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 1.3 
Nutze die `.` Syntax um jedes Element von `A` um 1 zu erhöhen und weise es einer Variable `A2` zu.

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

## Übung: Convert

#### 2.1 
Schau dir die Dokumentation der `convert` Funktion an. Mit `?` kannst du die hilfe/dokumentation der nachfolgenden Eingabe ausgeben lassen. 

```julia

```

#### 2.2 
Lege eine Variable mit dem wert `365` und dem Namen `days` an. Konvertiere days zu float und weiße sie der Variablen `days_float` zu. 

```julia

```

```julia
@assert days == 365
@assert days_float == 365.0

```

#### 2.3 
Führe die folgenden Zeilen aus und schau was passiert. 

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia

```