# Multiple dispatch

In diesem Abschnitt schauen wir uns **multiple dispatch** an, welches ein Schlüsseleigenschaft von Julia ist.

#### Was wir schon kennen

Wir können Funktionen in Julia deklarieren ohne Julia irgendwelche Informationen über die Typen der Eingabeparameter zu geben. Julia wird die Typen der Eingabeparameter selbst herausfinden. 
```julia
function f(x)
    return x + 1
end
```

```julia
f(x) = x.^2
f(10)
```

Deshalb können wir der Funktion auch unterschiedliche Datentypen Übergeben:

```julia
f([1, 2, 3])
```

#### Typ der Eingabeparameter explizit angeben

Alternativ können wir Julia auch sagen, welche Typen die Eingabeparameter haben dürfen. 

Zum Beispiel schreiben wir eine Funktion `foo`, die nur Strings als Eingabeparameter akzeptiert.

```julia
foo(x::String, y::String) = println("My inputs x and y are both strings!")
```

Dem eingabeargument folgen zei Doppelpunkte und dann der Typ des Eingabearguments.

Jetzt funktioniert `foo` nur noch mit Strings:

```julia
foo("hello", "hi!")
```

```julia
foo(3, 4)
```

Um die Funktion `foo` mit anderen Datentypen zu verwenden, müssen wir die Funktion erneut mit den gewünschten Datentypen deklarieren.

```julia
foo(x::Int, y::Int) = println("My inputs x and y are both integers!")
```

```julia
foo(3, 4)
```

Jetzt funktioniert `foo` mit Integern und mit Strings:

```julia
foo("hello", "hi!")
```

Das bringt uns langsam zum Kern von Multiple Dispatch. Als wir deklarierten

```julia
foo(x::Int, y::Int) = println("My inputs x and y are both integers!")
```
haben damit aber nicht 
```julia
foo(x::String, y::String)
```
überschrieben oder ersetzt. Stattdessen haben wir lediglich eine zusätzliche ***Methode*** zur ***generischen Funktion*** namens `foo` hinzugefügt.


Eine ***generische Funktion***  ist das abstrakte Konzept, das mit einer bestimmten Operation verbunden ist.

Zum Beispiel repräsentiert die generische Funktion `+` das Konzept der Addition.

Eine ***Methode*** st eine spezifische Implementierung einer generischen Funktion für bestimmte Argumenttypen.

Zum Beispiel hat `+` Methoden, die Fließkommazahlen, Ganzzahlen, Matrizen usw. akzeptieren.

Wir können `methods` verwenden, um zu sehen, wie viele Methoden es für `foo` gibt.

```julia
methods(foo)
```

Das Funktioniert natürlich auch für bereits bekannte Funktionen wie `+`:

```julia
methods(+)
```

Wir können `foo` also sowohl für Ganzzahlen als auch für Zeichenketten aufrufen. Wenn `foo` mit einem bestimmten Satz von Argumenten aufgerufen wird, wird Julia die Typen der Eingaben ableiten und die entsprechende Methode aufrufen. Das ist Multiple Dispatch.

Multiple Dispatch macht unseren Code generisch und schnell. Unser Code kann generisch und flexibel sein, weil wir Code in Bezug auf abstrakte Operationen wie Addition und Multiplikation schreiben können, anstatt in Bezug auf spezifische Implementierungen. Gleichzeitig läuft unser Code schnell, weil Julia in der Lage ist, effiziente Methoden für die relevanten Typen aufzurufen.

Um zu sehen, welche Methode aufgerufen wird, wenn wir eine generische Funktion aufrufen, können wir das `@which`-Makro verwenden:

```julia
@which foo(3, 4)
```

Der `@which` Operator für die Funktion `+` gibt uns die Methode an, die aufgerufen wird, wenn wir:

```julia
@which 3.0 + 3.0
```

Es gibt auch abstrakte Typen, die wir verwenden können, um Methoden zu deklarieren. Zum Beispiel können wir eine Methode für `foo` deklarieren, die nur ***subtypes*** von `Number` akzeptiert (also `Int`, `Float64`, usw.). Wir können dies tun, indem wir den Typ `Number` als Eingabeargument angeben:

```julia
foo(x::Number, y::Number) = println("My inputs x and y are both numbers!")
```

Jetzt funkioniert `foo` auch mit `Number`-subtypes:

```julia
foo(3.0, 4.0)
```

Wir können zudem eine fallback Methode für `foo` deklarieren, die für alle anderen Datentypen funktioniert:

```julia
foo(x, y) = println("I accept inputs of any type!")
```

```julia
v = rand(3)
foo(v, v)
```

### Übung: Multiple Dispatch

#### 4.1

Erweitere die Funktion `foo`, indem du eine Methode hinzufügst, die nur einen Eingabeargument akzeptiert, welcher vom Typ `Bool` ist und "foo with one boolean!" zurückgibt.

```julia

```

#### 4.2

Prüfe ob du die Methode richtig deklariert hast, indem du die Funktion `foo` mit einem `Bool`-Eingabeargument aufrufst.

```julia

```

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert foo(true) == "foo with one boolean!"
```